# Exercise 2

## physics718: Programming in Physics and Astronomy with C++ or Python (SS 2020)
Oliver Cordes & Thomas Erben

Homework is due on **Tuesday, 05/05/2020, 23:59pm**

 * You only learn a programming language by actively praticing and using it! We therefore **strongly** advise you to work on the homework problems.
 * Please discuss the problems with your student peers in case of problems.
 * Your code(s) needs to be well and appropriately commented!
 * All files of your solution should be in the same directory as this notebook
 * Finally submit this notebook with your solutions in the nbgrader system
 
**Main topics of this exercise:**
 * work with scalar types in Python *int*, *float* and *bool*
 * control structures *if* and *while*
 * floating poing accuracy

Put in here your group number: 

Group 26

# 1. Numbers to binaries (20 points)

## 1.1) Integers to binaries

As mentioned in the lecture integer numbers can be written in different number systems. For us humans we are working in the decimal system but the computers are dealing with numbers in the binary system. Theoretically all integer numbers $n$ can be written in number systems with a base $b$:

$$ n = \sum_{i=0}^{m} q_{i}b^{i}$$

with $0 \le q_{i} < b$. As said we humans use $b=10$ and the computer $b=2$. 

In this first task we want to convert numbers into the binary system. Here are some examples:
* $1_{10}$ -> $1_2$
* $3_{10}$ -> $11_2$
* $14_{10}$ -> $1110_2$
etc.

For the conversion you can use the following algorithm. Start with the number to be converted and divide this number by 2 and store the quotient and the rest. Then proceed with the quotient and divide this number again until the quotient reaches 0. The rest numbers written in the reverse order gives the representation of $n$ in the binary system.

Here is an example which converts $n=14$ into the binary system:
$$ 14 :2 = 7 + 0$$ 
the quotient is 7  and the rest is 0, then:
$$ 7 :2 = 3 + 1 $$
$$ 3 :2 = 1 + 1 $$
$$ 1 :2 = 0 + 1 $$

The rest numbers written in the reverse order are now $1110$ which is 14 written in the binary system!

**Your task:**

Implement the algorithm in a single notebook code cell. Start with a predefined variable $n$ which contains the number to convert. The converted number should be stored in a string with the name *output_integer*. Print this string at the end of the code. The value of *output_integer* will be tested in the grading system, so for the validation please convert the $n=14$!

**Hints:**
 * on this [webpage](https://indepth.dev/the-simple-math-behind-decimal-binary-conversion-algorithms/) you can see how and why the algorithm is working (this page uses the base 2 and describes the decimal to binary conversion)
 * you can do all the divisions in a while loop
 * for calculating the quotient and the rest you can use the [*divmod*](https://docs.python.org/3/library/functions.html#divmod)-command in python.
   
   Example:
   ```Python
   quot, rest = divmod(a,b)
   ```
 * for the representation of the converted number you can start with an empty string e.g. ```s = ''``` and prepend a rest
   number in every step with ```s = str(rest) + s``` !
 * test your code with different examples, e.g. $n=14$ and $n=0$!

<!-- grade_id: cell-249cdf3b0452f1ac, points: 0, grade: false, solution: true -->
```python
b = 2

# YOUR CODE HERE

def convert2(n,b):                           
    n_2 = ''                                 # Initialize string to be returned and set first value for quot
    quot = n
    if n==0:                                 # check whether n=0 before going into the while loop
        return'0' 
    while quot!=0:                           # repeat loop until quot is zero
        quot, rest = divmod(quot,b)          # for conversion into the binary system, choose  division by two
        n_2 = str(rest) + n_2                # add element to existing string  
    return n_2


convert2(14,b)
convert2(63,b)
```

Output:
```
'111111'
```

<!-- grade_id: cell-bcf437cbac810793, points: 10, grade: true, solution: false -->
```python
assert(output_integer == '1110')
```

<!-- grade_id: cell-bcf437cbac810793, points: 10, grade: true, solution: false -->
```python
assert(output_integer == '1110')
```

## 1.2) Floats to binaries

Converting floating point number into the binary system is similar to the integer conversion. The sum as shown above can be extended with negative indices. 

$$ x = \sum_{i=-k}^{m} q_{i}b^{i}$$

Here are a few examples:
 * $1.5_{10}$  -> $1.1_2$
 * $3.25_{10}$ -> $11.01_2$
 ...
 
 
For the conversion of a float number $x$ one can use the following algorithm. First we will multiply $x$ with 2 until $x$ is (approximately) an integer. We should store the number of multiplications. Then we can convert this int into a binary using the algorithm which you implemented above. To convert the binary into a binary float we can add a *binary point* at the end and shift the point to the left by the number of multiplication positions.

Let's convert the floating point number $x=3.25$:

1. Do the multiplications: $3.25 * 2 = 6.5$ -> $6.5 * 2 = 13$ -> which means 2 multiplications until we have an integer
2. convert 13 into a binary -> $1101$
3. add a binary point at the end: $1101.$
3. shift the point by 2 positions to the left, to compensate for the multiplications -> $x_b = 11.01_2$

A second example, convert $x=0.125$:

1. Do the multiplications: $0.125 * 2 * 2 * 2 = 1$ -> 3 multiplications
2. convert 1 into a binary (trivial) -> $1$
3. now we need to prepend $n*0$ (n zeros, n = min. the number of multiplications) and the binary point at the end: $0001.$
4. shift the point by 3 positions to the left, to compensate for the multiplications -> $x_b = 0.001_2$

**Your task**:

Implement this algorithm. You can use the code from the previous task to convert an integer into a binary. 
For some of the floating point numbers there will be not an exact binary representation. You can stop the multiplication after you calculated a value which differs less than $10^{-6}$ from its integer representation. 
If you have done more than 10 multiplications you should print a warning at the end of your code indicating that the binary representation is probably not finite.

**Hints**:
 * to convert an floating point number $x$ to an integer you can use ```int(x)```! ```int(x)``` truncates $x$ towards zero.
 * use the function ```abs(x-int(x))``` to test the accuracy of the interger conversion
 * to shift the binary point, you can directly put the binary point at the correct place (don't add the point at the end before).    For adding the point, please  have a look at the following cell with a small introduction into *slices* of array and array like structures e.g. strings

```python
# strings are represented similar to arrays
s = '123456789'

# each character can accessed by their position
print(s[0])   # the first character, starting with 0 until length of the string -1
print(s[-1])  # negative indices counts from the end

# you can cut a smaller string out of the larger string, by defining start and end-indices
print(s[1:3]) # print the second and the third character
print(s[1:])  # if you leave the second limit open, it means until the end
print(s[:2])  # if you leave the beginning open, it means from the beginning
print(s[-2:]) # print the last two characters 

point = 6
# if you want to place a point between the the first  and the last 5 characters
# you need to cut the string into 2 pieces and put the '.' in between
s = s[:-point] + '.' + s[-point:]
print(s)
```

Output:
```
1
9
23
23456789
12
89
123.456789
```

<!-- grade_id: cell-95b6c4063d5857c1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def convert_float(f,b):
    mult_f = f                                                   # stores value after multiplications
    n_mult = 0                                                   # stores number of multiplications
    epsilon = 10**(-6)
    while abs(mult_f -int(mult_f)) > epsilon:                    # multiplies by two until precision is reached
        n_mult += 1
        mult_f = 2*mult_f
        
    s = convert2(int(mult_f),b)                                    # original algorithm with int argument
    if n_mult != 0:                                              # adding the point in case f was not an integer                                           
        
              
        if f<1:                                                  # f<1 requires adding further zeros to the depending on the number of multiplications
            for i in range(n_mult):
                
                s = str(0) + s 
        s = s[:-n_mult] + '.' + s[-n_mult:]
    if (n_mult)>=10 :
        print('The binary representation is probably not finite.') 
    return(s)


#convert_float(14)
#convert_float(0)
#convert_float(0.5)
#convert_float(0.25)
#convert_float(1.5)
convert_float(1.25)
convert_float(1.2358765695)
```

Output:
```
The binary representation is probably not finite.
'1.0011110001100010011010000010011111100101001010011111'
```

<!-- grade_id: cell-ab47346348a927a3, points: 10, grade: true, solution: false -->
```python
assert(output_float.lstrip('0') == '11.01')
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-69-1801d02375d8> in <module>
----> 1 assert(output_float.lstrip('0') == '11.01')

NameError: name 'output_float' is not defined
```

# 2. Problems with an integral series (15 points)

Consider the sequence
$$
E_n=\int_0^1x^n\exp(x-1)\,\mathrm{d}x; \qquad n=1,2,\dots
$$
One can show that it converges to zero: $\lim_{n\to\infty}E_n=0$. Integrating by parts, we obtain the following relation between two consecutive elements:
$$
E_n = 1-nE_{n-1} \text{ with } E_1=1-\int_0^1\exp(x-1)\,\mathrm{d}x=\exp(-1).
$$
We want to estimate $E_{20}$. Do this by computing and printing the first 20 elements of the series. Obtain a second estimate of $E_{20}$ by reverting the relation:
$$
  E_{n-1} = \frac 1n (1-E_{n}) \text{ with } E_{50} = 0
$$
and printing the elements $E_{50}, E_{49}, \dots, E_{20}$.

**Your tasks:**

Create a cell for each experiment/program and a markdown cell in which you argue which one of the results you trust more and try to explain why they disagree that much.


**Hints:** 
 * Assume that in the first case, $E_1$ is represented internally as a float number with an error, i.e. $E_1 = E^t_1 + \epsilon$, where $E^t_1$ is the true value and $\epsilon$ the error. We know that $\epsilon\approx 10^{-18}$ for `Python` float numbers. What happens with $\epsilon$ when you calculate new elements of the series? 
 * for the exponential function you can use the numpy module with ```import numpy``` and use the defined function ```numpy.exp(x)``` to calculate $e^x$!

<!-- grade_id: cell-3eae773e632eab70, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
import numpy as np

E=np.zeros(50)           #create an array with 50 zeros
E[0]=np.exp(-1)          #initialize E_1
for i in range(1,20):    #iterate till E_20
    E[i]=1-i*E[i-1]
    
print(E[0:20])           #print the first 20 elements
```

Output:
```
[ 3.67879441e-01  6.32120559e-01 -2.64241118e-01  1.79272335e+00
 -6.17089341e+00  3.18544671e+01 -1.90126802e+02  1.33188762e+03
 -1.06541009e+04  9.58879084e+04 -9.58878084e+05  1.05476599e+07
 -1.26571918e+08  1.64543494e+09 -2.30360891e+10  3.45541337e+11
 -5.52866138e+12  9.39872435e+13 -1.69177038e+15  3.21436373e+16]
```

<!-- grade_id: cell-8ec1cc6f8ddca053, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
for j in range(1,31):    #back iteration til E_20
    i=50-j
    E[i-1]=(1-E[i])/i
    
print(E[-31:])           #print E_20, E_21,..., E_50

```

Output:
```
[0.04772276 0.04554488 0.04355743 0.04173644 0.04006181 0.03851655
 0.03708621 0.03575842 0.03452253 0.03336929 0.03229068 0.03127967
 0.03033011 0.02943654 0.02859416 0.02779868 0.02704629 0.02633358
 0.02565749 0.02501527 0.02440443 0.02382273 0.02326812 0.02273877
 0.02223297 0.0217492  0.02128604 0.02084238 0.02040816 0.02040816
 0.        ]
```

<!-- grade_id: cell-9e84b36eaf2530ef, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

The second method is more believable. 

The first method magnifies the inaccuracy every time by $n$. By the 20th element the inaccuracy has reached $20!\times 10^{-18}\approx 2$. More importantly, the oscillation of positive and negative results in the first method is a clear indicator of mistakes because the integral defined above is always positive. This happens because of the nature of the recurrsion relation $E_n = 1-nE_{n-1}$: magnification by $n$ makes the next element sensitive to sign switching.


```python

```